---
id: 6735a7370e0ae93a4577c689
title: Design a Newspaper Layout
challengeType: 25
dashedName: design-a-newspaper-layout
demoType: onClick
---

# --description--

在这个实验中，您将使用CSS Grid设计一个报纸布局，包括网格行和网格列等概念。

**目标：** 实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 您的页面应该包含一个带有[newspaper-layout](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L340-L355)类的[main](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L280-L326)元素。
2. .newspaper-layout应该包含一个带有[title](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L357-L363)类的[header](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L281-L283)，其中包含一个显示报纸名称的h1元素。
3. .newspaper-layout应该包含一个带有[feature-article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L378-L380)类的[article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L285-L293)元素，用于主要新闻文章。
4. .feature-article应该包含一个用于文章标题的h2元素，后跟一个用于文章内容的p元素。
5. 您应该再添加三个[article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L285-L293)元素用于较小的文章，类名分别为[small-article1](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L394-L396)、[small-article2](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L398-L400)和[small-article3](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L402-L404)。
6. 每个较小的文章都应该包含一个用于文章标题的h3元素，后跟一个用于文章内容的p元素。
7. .newspaper-layout应该包含一个带有[secondary-article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L382-L384)类的[article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L285-L293)元素，用于额外的新闻部分。
8. .secondary-article应该包含一个用于文章标题的h2元素，后跟一个用于文章内容的p元素。
9. .newspaper-layout应该包含一个带有[cover-image](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L406-L414)类的[figure](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L323-L325)元素，用于显示代表报纸内容的图像。
10. 带有[title](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L357-L363)、[feature-article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L378-L380)、[secondary-article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L382-L384)、[cover-image](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L406-L414)、[small-article1](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L394-L396)、[small-article2](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L398-L400)和[small-article3](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L402-L404)类的元素应该有一个设置为相同类名的[grid-area](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L358-L358)属性。
11. 您的.newspaper-layout应该使用CSS Grid和[grid-template-areas](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L342-L346)属性来定义各部分的排列：
    - .title应该跨越顶部行。
    - .feature-article应该在第二行跨越两列，.cover-image在第三列。
    - .secondary-article应该在第三行跨越两列，.cover-image在第三列。
    - 三个小文章应该出现在第四行。
12. .newspaper-layout应该有一个[grid-template-columns](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L347-L347)属性，将空间分为三等列。
13. 您应该将.newspaper-layout的[grid-template-rows](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L348-L348)属性设置为第一行为[auto](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L207-L225)，其余行的空间平均分配。
14. 您应该在网格项目之间添加间隙。
15. 通过将.max-width设置为[100%](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-celestial-bodies-database/66db559d34c7089b9b41bfd5.md#L47-L65)，确保.cover-image内的图像适合指定空间。

**注意：** 确保在HTML中链接样式表并应用CSS。

# --hints--

您应该有一个[main](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L280-L326)元素。

```js
assert.exists(document.querySelector('main'));
```

您的[main](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L280-L326)元素应该有一个[newspaper-layout](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L340-L355)的[class](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L50-L51)。

```js
assert.exists(document.querySelector('main.newspaper-layout'));
```

您应该在.newspaper-layout元素内有一个带有[title](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L357-L363)类的[header](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L281-L283)元素。

```js
const header = document.querySelector(".newspaper-layout > header");
assert.exists(header);
assert(header.classList.contains("title"));
```

您的[header](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L281-L283)元素应该包含一个h1元素。

```js
const h1 = document.querySelector("header h1");
assert.exists(h1);
```

您的.title元素应该包含一些文本。

```js
assert.isAbove(document.querySelector(".title").innerText.length, 0);
```

您应该有五个[article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L285-L293)元素，类名分别为[feature-article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L378-L380)、[small-article1](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L394-L396)、[small-article2](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L398-L400)、[small-article3](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L402-L404)和[secondary-article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L382-L384)。

```js
const featureArticle = document.querySelector('article.feature-article');
const smallArticle1 = document.querySelector('article.small-article1');
const smallArticle2 = document.querySelector('article.small-article2');
const smallArticle3 = document.querySelector('article.small-article3');
const secondaryArticle = document.querySelector('article.secondary-article');

assert.exists(featureArticle);
assert.exists(smallArticle1);
assert.exists(smallArticle2);
assert.exists(smallArticle3);
assert.exists(secondaryArticle);
```

.feature-article应该包含一个h2元素，后跟一个p元素。

```js
const h2 = document.querySelector(".feature-article h2");
const p = document.querySelector(".feature-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

.small-article1、.small-article2和.small-article3应该包含一个h3元素，后跟一个p元素。

```js
for (let i = 1; i < 4; i++) {
    let h3 = document.querySelector(`.small-article${i} h3`);
    let p = document.querySelector(`.small-article${i} h3 + p`);
    assert.exists(h3);
    assert.exists(p);
}
```

.secondary-article应该包含一个h2元素，后跟一个p元素。

```js
const h2 = document.querySelector(".secondary-article h2");
const p = document.querySelector(".secondary-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

您应该在.newspaper-layout内有一个带有[cover-image](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L406-L414)类的[figure](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L323-L325)元素。

```js
assert.exists(document.querySelector(".newspaper-layout figure.cover-image"));
```

您的.cover-image应该包含一个img元素。

```js
assert.exists(document.querySelector(".cover-image img"));
```

.newspaper-layout元素的[display](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L134-L134)属性应该设置为[grid](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L341-L341)。

```js
assert.equal(getComputedStyle(document.querySelector('.newspaper-layout')).display, 'grid');
```

您的.newspaper-layout元素应该有一个[grid-template-areas](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L342-L346)属性，其中[title](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L357-L363)跨越第一行的所有三列。

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
assert.match(newspaperLayout.gridTemplateAreas, /^("|')title title title\1/);
```

网格模板的第二行应该包含跨越两列的[feature-article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L378-L380)和[cover-image](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L406-L414)。

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromSecondRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").trimStart();
assert.match(fromSecondRow, /^("|')feature-article feature-article cover-image\1/);
```

网格模板的第三行应该包含跨越两列的[secondary-article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L382-L384)和[cover-image](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L406-L414)。

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromThirdRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromThirdRow, /^("|')secondary-article secondary-article cover-image\1/);
```

网格模板的第四行应该包含[small-article1](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L394-L396)、[small-article2](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L398-L400)和[small-article3](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L402-L404)。

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromFourthRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromFourthRow, /^("|')small-article1 small-article2 small-article3\1$/);
```

您的.newspaper-layout选择器应该使用[grid-template-columns](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L347-L347)属性将空间分为三个相等的部分。记住您可以使用[repeat](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L174-L177)函数或手动定义每个分数单位。

```js
assert.oneOf(
    new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateColumns,
    ['1fr 1fr 1fr', 'repeat(3, 1fr)'] 
);
```

您的.newspaper-layout选择器应该有一个设置为[auto 1fr 1fr 1fr](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L183-L183)的[grid-template-rows](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L348-L348)属性。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateRows, ['auto 1fr 1fr 1fr', 'auto repeat(3, 1fr)']);
```

您应该在网格项目之间添加间隙。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gap);
```

.title应该有一个设置为[title](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L357-L363)的[grid-area](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L358-L358)。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.title')?.gridArea, 'title');
```

.feature-article应该有一个设置为[feature-article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L378-L380)的[grid-area](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L198-L202)。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.feature-article')?.gridArea, 'feature-article');
```

.cover-image应该有一个设置为[cover-image](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L406-L414)的[grid-area](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L204-L208)。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cover-image')?.gridArea, 'cover-image');
```

.secondary-article应该有一个设置为[secondary-article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L382-L384)的[grid-area](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L210-L214)。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.secondary-article')?.gridArea, 'secondary-article');
```

.small-article1应该有一个设置为[small-article1](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L394-L396)的[grid-area](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L216-L220)。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article1')?.gridArea, 'small-article1');
```

.small-article2应该有一个设置为[small-article2](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L398-L400)的[grid-area](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L222-L226)。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article2')?.gridArea, 'small-article2');
```

.small-article3应该有一个设置为[small-article3](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L402-L404)的[grid-area](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L228-L232)。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article3')?.gridArea, 'small-article3');
```

.cover-image内的img应该有一个[max-width](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-newspaper-layout/6735a7370e0ae93a4577c689.md#L417-L417)为[100%](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-celestial-bodies-database/66db559d34c7089b9b41bfd5.md#L47-L65)。

```js
const coverImage = document.querySelector(".cover-image");
const img = document.querySelector(".cover-image img");
assert.equal(getComputedStyle(img).maxWidth, "100%");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
<meta charset="utf-8">
<title>Newspaper Layout</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Layout</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main class="newspaper-layout">
        <header class="title">
            <h1>The Daily Local News</h1>
        </header>

        <article class="feature-article">
            <h2>Breaking News: Volcano Eruption Disrupts Tourism</h2>
            <p>Recently, a volcano erupted in a popular tourist destination. The eruption has caused widespread panic
                and has disrupted tourism in the area. The volcano has been spewing lava and ash for several days now,
                and authorities are urging residents and tourists to evacuate the area immediately. The eruption has
                also caused a number of flights to be cancelled, leaving many tourists stranded. The situation is still
                developing, and authorities are working to contain the eruption and ensure the safety of everyone in the
                area.</p>
        </article>

        <article class="small-article1">
            <h3>Sports: Local Team Wins Championship</h3>
            <p>Hockey fans are celebrating today as the local team has won the championship. The team, which has been on
                a winning streak all season, clinched the title in a thrilling final match. Fans took to the streets to
                celebrate the victory, waving flags and chanting the team's name.</p>
        </article>

        <article class="small-article2">
            <h3>Health: Tips for a Balanced Diet</h3>
            <p>A diet high in calories, sugar, and unhealthy fats can lead to a variety of health problems, including
                obesity, diabetes, and heart disease. To maintain a healthy weight and reduce your risk of chronic
                diseases, it's important to eat a balanced diet that includes a variety of nutrient-rich foods. Here are
                some tips for eating a balanced diet:</p>
        </article>

        <article class="small-article3">
            <h3>Travel: Top 10 Destinations for 2025</h3>
            <p>Traveling is one of the best ways to experience new cultures, meet new people, and see the world. If
                you're looking for inspiration for your next trip, here are the top 10 destinations for 2025:</p>
        </article>

        <article class="secondary-article">
            <h2>Technology: The Rise of AI</h2>
            <p>Artificial intelligence (AI) is changing the way we live and work. From self-driving cars to virtual
                assistants, AI is revolutionizing the world around us. But what exactly is AI, and how does it work? In
                this article, we'll explore the rise of AI and its impact on society.</p>
        </article>

        <figure class="cover-image">
            <img src="https://cdn.MasterPuti.org/curriculum/labs/volcano.jpg" alt="Cover Image">
        </figure>
    </main>
</body>

</html>
```

```css
body {
    font-family: 'Georgia', 'Times New Roman', Times, serif;
    padding: 0;
    margin: 0;
    background-color: #f0f0f0;
}

.newspaper-layout {
    display: grid;
    grid-template-areas:
        "title title title"
        "feature-article feature-article cover-image"
        "secondary-article secondary-article cover-image"
        "small-article1 small-article2 small-article3";
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: auto 1fr 1fr 1fr;
    gap: 15px;
    background-color: #fff;
    padding: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    box-sizing: border-box;
    text-align: justify;
}

.title {
    grid-area: title;
    background-color: #1B1B32;
    color: white;
    padding: 15px;
    text-align: center;
}

.title h1 {
    font-size: 2.5rem;
    line-height: 1.3;
    margin: 0;
}

.feature-article,
.secondary-article {
    background-color: #f9f9f9;
    padding: 15px;
    font-size: 1rem;
}

.feature-article {
    grid-area: feature-article;
}

.secondary-article {
    grid-area: secondary-article;
}

.small-article1,
.small-article2,
.small-article3 {
    background-color: #e9e9e9;
    padding: 15px;
    font-size: 0.9rem;
}

.small-article1 {
    grid-area: small-article1;
}

.small-article2 {
    grid-area: small-article2;
}

.small-article3 {
    grid-area: small-article3;
}

.cover-image {
    grid-area: cover-image;
    display: flex;
    justify-content: center;
    align-items: center;
    background-color: #ddd;
    padding: 15px;
    margin: 0;
}

.cover-image img {
    max-width: 100%;
    height: 350px;
    width: 350px;
    border-radius: 5px;
}
```

