---
id: 66b59829dba144ff1351220f
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

接下来，我们用 `botName` 变量来组合一句问候语吧！使用模板字面量把 `botName` 嵌入到字符串中，创建变量 `greeting`，然后把它打印到控制台。做得好，你很棒！

# --hints--

小提示：

- 创建 `greeting` 变量并赋值为包含 `botName` 的字符串。

```js
assert.isNotNull(greeting);
```

它应为字符串：

```js
assert.isString(greeting);
```

并且应等于 `"My name is teacherBot."`：

```js
assert.equal(greeting, "My name is teacherBot.");
```

别忘了用 `console.log` 打印 `greeting`。

```js
assert.match(code, /console\.log\(.*greeting.*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
